## Exercise 06.1 (selecting and passing data structures)

The task in Exercise 04 for computing the area of a triangle involved a function with six arguments ($x$ and $y$ components of each vertex). With six arguments, the likelihood of a user passing arguments in the wrong order is high. 

Use an appropriate data structure, e.g. a `list`, `tuple`, `dict`, etc,  to develop a new version of the function with a simpler interface (the interface is the arguments that are passed to the function). Add appropriate checks inside your function to validate the input data.

<!-- grade_id: cell-50c4566d7bcbffaf, points: 0, grade: false, solution: true -->
```python
def area(coords):
  verifFlag = True
  if (len(coords) == 3):
    for x in coords:
      if (len(x) == 2):
        for n in x:
          if isinstance(n, int) or isinstance(n, float):
            continue
          else:
            print("Error: non integer tuples")
            return None
      else:
        print("Error: non 2D co-ords")
        return None
  else:
    print("Error: please enter 3 co-ords")
    return None
  
  area = abs(coords[0][0] * (coords[1][1]-coords[2][1]) + 
            coords[1][0] * (coords[2][1]-coords[0][1]) +
            coords[2][0] * (coords[0][1]-coords[1][1]))/2
  return area
  return None


coords = [(0, 0), (1, 0), (0, 1)] #0, 1, 2
area(coords)

```

Output:
```
Error: more than 3 co-ords
```

## Exercise 06.2 (selecting data structures)

For a simple (non-intersecting) polygon with $n$ vertices, $(x_0, y_0)$, $(x_1, y_1)$, . . , $(x_{n-1}, y_{n-1})$, the area $A$ is given by
$$
A = \left| \frac{1}{2} \sum_{i=0}^{n-1} \left(x_{i} y _{i+1} - x_{i+1} y_{i} \right) \right|
$$
and where $(x_n, y_n) = (x_0, y_0)$. The vertices should be ordered as you move around the polygon.

Write a function that computes the area of a simple polygon with an arbitrary number of vertices. Test your function for some simple shapes. Pay careful attention to the range of any loops.

<!-- grade_id: cell-4b5a619405787746, points: 0, grade: false, solution: true -->
```python

def polyarea(v):
  v.append(v[0])  # set (x_n, y_n) = (x_0, y_0)
  sum = 0
  for i in range(len(v)-1):
    sum += (v[i][0]*v[i+1][1] - v[i+1][0]*v[i][1])
    print(i)
  area = abs(sum/2)
  return area


coords = [(1, 1), (1, 2), (2, 2), (2, 1)]
print(polyarea(coords))

```

Output:
```
0
1
2
3
1.0
```

## Exercise 06.3 (indexing)

Write a function that uses list indexing to add two vectors of arbitrary length, and returns the new vector. Include a check that the vector sizes match, and print a warning message if there is a size mismatch. The more error information you provide, the easier it would be for someone using your function to debug their code.

Add some tests of your code.

#### Hint: You can create a list of zeros of length `n` by

    z = [0]*n
    
#### Optional (advanced) 

Try writing a one-line version of this operation using list comprehension and the built-in function [`zip`](https://docs.python.org/3/library/functions.html#zip).

<!-- grade_id: cell-47daca03f7a96b62, points: 0, grade: false, solution: true -->
```python
def sum_vector(x, y):
    "Return sum of two vectors"
    if len(a) == len(b):
        tot = [0]*len(a)
    else:
        print("Error: Vector size mismatch")
        return None
    tot = [a+b for a,b in zip(x,y)]
    return tot

a = [0, 2, -5, 7]
b = [6, 7, 3, 5]
c = sum_vector(a, b)
print(c)

a = [3.2, 8.2, 0, 7]
b = [-2, 7, -15, 1]
c = sum_vector(a, b)
print(c)
```

Output:
```
[6, 9, -2, 12]
[1.2000000000000002, 15.2, -15, 8]
```

<!-- grade_id: cell-95f5819810d2b2c7, points: 0, grade: true, solution: false -->
```python
a = [0, 4.3, -5, 7]
b = [-2, 7, -15, 1]
c = sum_vector(a, b)
assert c == [-2, 11.3, -20, 8]
```

### Extension: list comprehension

<!-- grade_id: cell-c182b3f0d8952692, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Exercise 06.4 (dictionaries)

Create a dictionary that maps college names (the key) to college abbreviations for at least 5 colleges
(you can find abbreviations at https://en.wikipedia.org/wiki/Colleges_of_the_University_of_Cambridge#Colleges).
From the dictionary, produce and print

1. A dictionary from college abbreviation to name; and
1. A list of college abbreviations sorted into alphabetical order.

*Optional extension:* Create a dictionary that maps college names (the key) to dictionaries of:

- College abbreviation
- Year of foundation 
- Total number students
 
for at least 5 colleges. Take the data from https://en.wikipedia.org/wiki/Colleges_of_the_University_of_Cambridge#Colleges. Using this dictionary, 

1. Find the college with the greatest number of students and print the abbreviation; and 
2. Find the oldest college, and print the number of students and the abbreviation for this college.

<!-- grade_id: cell-bc271dc3885e19c5, points: 0, grade: false, solution: true -->
```python
coltoabv = {
  "Homerton": "HO",
  "Jesus": "JE",
  "Selwyn": "SE",
  "Peterhouse" : "PET",
  "Trinity Hall": "TH"
} 

abvtocol = {}
for n, m in coltoabv.items():
  abvtocol[m] = n #create new entries in flipped dict with the key as the ABVs

sortedabvs = sorted(list(abvtocol))

print(abvtocol)
print(sortedabvs)


```

Output:
```
{'HO': 'Homerton', 'JE': 'Jesus', 'SE': 'Selwyn', 'PET': 'Peterhouse', 'TH': 'Trinity Hall'}
['HO', 'JE', 'PET', 'SE', 'TH']
```

#### Optional extension

<!-- grade_id: cell-7040c4053c72ce09, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```